# Milestone 2: Preprocessing

## 1. Import Libraries

```python
from collections import Counter
import math
import time

import matplotlib as mpl
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
import seaborn as sns
from scipy import stats
from sklearn.linear_model import Lasso, LassoCV
from sklearn.linear_model import LinearRegression
from sklearn.metrics import mean_squared_error
from sklearn.model_selection import cross_validate
from sklearn.model_selection import train_test_split
from sklearn.preprocessing import PolynomialFeatures
from sklearn.utils import shuffle

from warnings import simplefilter
simplefilter('ignore', category=FutureWarning)
```

## 2. Access the data that you will be using for the final project by downloading, collecting, or scraping* from the relevant source(s)

We accessed **fertility rate by state over time** (the response variable) from https://www.cdc.gov/nchs/pressroom/sosmap/fertility_rate/fertility_rates.htm and saved it as `fertility_rate_census.csv`

We downloaded each of the following datasets to be a possible predictor variable

**Age at first marriage** from _ saved as `name`

## 3. Load the data into a Jupyter notebook and understand the data by examining, among other characteristics of interest, data missingness, imbalance, and scaling issues.

```python
fertility_rate_by_state_by_time = pd.read_csv('data/fertility_rate_census.csv')
fertility_rate_by_state_by_time.head()
```

Output:
```
YEAR STATE  FERTILITY RATE  BIRTHS  \
0  2022    AL            58.7   58149   
1  2022    AK            64.9    9359   
2  2022    AZ            54.9   78547   
3  2022    AR            60.2   35471   
4  2022    CA            52.8  419104   

                                        URL  
0     /nchs/pressroom/states/alabama/al.htm  
1      /nchs/pressroom/states/alaska/ak.htm  
2     /nchs/pressroom/states/arizona/az.htm  
3    /nchs/pressroom/states/arkansas/ar.htm  
4  /nchs/pressroom/states/california/ca.htm
```

```python
fertility_rate_by_state_by_time.describe()
```

Output:
```
YEAR  FERTILITY RATE         BIRTHS
count   501.000000      501.000000     501.000000
mean   2016.710579       60.690419   76464.960080
std       4.611515        6.798393   87919.651528
min    2005.000000       44.300000    5133.000000
25%    2015.000000       56.400000   21903.000000
50%    2018.000000       60.600000   56170.000000
75%    2020.000000       64.300000   87193.000000
max    2022.000000       92.800000  548882.000000
```

## 4. Understand and describe the preprocessing required such that data is in a form amenable to later downstream tasks such as visualizing and modeling, as is appropriate to the specific project goals.

Tests for q1 (points: 10)
```
>>> expected_url = './data/html/screenboston.html'
>>> assert os.path.isfile(expected_url), f'Expected local file {expected_url}'
```
```
>>> expected_url = './data/html/screenboston.html'
>>> with open(expected_url, 'r') as f:
...     content = f.read()
...     s = BeautifulSoup(content)
...     assert s.select_one('p').text == 'Screen Boston', f'Content of file saved from {expected_url} should contain a <p> tag with the page name.'
```
```
>>>
```


Tests for q2 (points: 15)
```
>>> n = 50
>>> assert len(movies) > n, f'`movies` should contain more than {n} elements; you have {len(movies)}.'
>>> assert all((isinstance(m, dict) for m in movies)), 'Elements of `movies` should all be dictionaries.'
```
```
>>> keys = {'title', 'directors', 'year', 'genre', 'runtime', 'theater', 'screen_date', 'screen_times'}
>>> assert all((set(keys).issubset(m.keys()) for m in movies)), f'Each dictionary in `movies` should contain all of these keys: {', '.join(keys)}. At least one dictionary is missing one or more keys.'
```
```
>>> assert all((isinstance(m['year'], int) for m in movies)), "The 'year' value should be an integer in all dictionaries."
```


Tests for q3 (points: 15)
```
>>> snapshots = glob.glob('./data/html/snapshot_*.html')
>>> assert len(snapshots) >= 4, f'You should have at least 4 snapshots, but found {len(snapshots)} files with paths like ./data/html/snapshot_*.html'
```
```
>>> snapshots = glob.glob('./data/html/snapshot_*.html')
>>> assert all((re.match('snapshot_\\d{8}\\.html', os.path.basename(f)) for f in snapshots)), "All snapshot files should be named in the format 'snapshot_YYYYMMDD.html'"
```
```
>>> assert os.path.isfile('data/movies.json'), "The file 'data/movies.json' should exist."
```
```
>>> with open('data/movies.json', 'r') as f:
...     movies = json.load(f)
>>> n = 300
>>> assert len(movies) > n, f'`movies` should now contain more than {n} elements; you have {len(movies)}.'
```
```
>>> keys = {'title', 'directors', 'year', 'genre', 'runtime', 'theater', 'screen_date', 'screen_times'}
>>> assert all((isinstance(m, dict) and set(keys).issubset(m.keys()) for m in movies)), f'Each dictionary in `movies` should contain all of these keys: {', '.join(keys)}. At least one dictionary is missing one or more keys.'
```


Tests for q4 (points: 15)
```
>>> assert isinstance(df, pd.DataFrame), "You should have stored your DataFrame in a variable called 'df'"
>>> assert df.shape[1] == 8, 'Your DataFrame, df, should have 8 columns'
```
```
>>> assert df.shape[0] > 300, 'You should have found at least 300 non-duplicate rows'
>>> assert df.shape[0] == df.drop_duplicates().shape[0], 'There are still duplicate rows in your DataFrame'
```
```
>>> def is_datetime_column_hacky(column):
...     dtype_str = str(column.dtype)
...     valid_types = ['date', 'timestamp', 'timedelta', 'duration']
...     return any((t in dtype_str.lower() for t in valid_types))
>>> assert is_datetime_column_hacky(df['screen_date']), "The 'screen_date' column must be either a Pandas or PyArrow date/datetime type."
```
```
>>> def is_datetime_column_hacky(column):
...     dtype_str = str(column.dtype)
...     valid_types = ['date', 'timestamp', 'timedelta', 'duration']
...     return any((t in dtype_str.lower() for t in valid_types))
>>> assert is_datetime_column_hacky(df['runtime']), "The 'runtime' column must be either a Pandas or PyArrow timedelta type."
```
```
>>> assert df['screen_date'].is_monotonic_increasing, "The 'screen_date' column is not sorted in ascending order"
```
```
>>>
```


Tests for q6 (points: 7)
```
>>> assert 'wiki_id' in df.columns, "`df` should now have a column called 'wiki_id'."
```
```
>>> p = 0.5
>>> assert (q := df['wiki_id'].notna().mean()) >= p, f'You should have been able to find wiki IDs for at least {p:.0%} of the screenings. You found {q:.0%}.'
```
```
>>>
```


Tests for q7 (points: 3)
```
>>> assert 'wiki_html' in df.columns, "`df` should now have a column called 'wiki_html'."
```
```
>>> p = 0.5
>>> assert (q := df['wiki_html'].notna().mean()) >= p, f'You should have been able to acquire wiki page HTML content for at least {p:.0%} of the screenings. You found {q:.0%}.'
```


Tests for wrapup (points: 5)
```
>>> assert float(hours_spent_on_hw), 'Please select a time in hours (int or float) to specify how long you spent on this assignment.'
```
